# TP 1 : Manipuler python, numpy et matplotlib

## Exercice 1 : Quelques notions de base de python

1. Boucle *for* : Que font les codes suivants :

* Code 1 :
<p style="background:#F3F3F3;" >
<code style="background:#F3F3F3;">for i in [0,1,2,3]:
    print(i) </code> </p>  


* Code 2 : 
<p style="background:#F3F3F3;" >
<code style="background:#F3F3F3;">for j in range(4):
    print(j) </code> </p>


*Attention l'indentation peut s'avérer erronée lors de la copie, si le* <code> print </code> *n'apparait pas en vert corriger.* 
    
    

<!-- grade_id: cell-6b586ed23d37357c, points: 0, grade: false, solution: true -->
```python
# Question 1 :

print("Code 1 :\n")
# Code 1 
for i in [0,1,2,3]:
    print(i) 

print("\nCode 2 :\n")
# Code 2
for j in range(4):
    print(j)

# Ils font la même chose.
```

2. Condition *if* : Que fait le code suivant:

<p style="background:#F3F3F3">
<code style="background:#F3F3F3;"># Code à recopier
a= 8
if a>4:
    print("a est plus grand que 4")
else:
    print("a est plus petit ou egal a 4")</code> </p>

<!-- grade_id: cell-dee616af522a520a, points: 0, grade: false, solution: true -->
```python
# Question 2

a = 8
if a>4:
    print("a est plus grand que 4")
else:
    print("a est plus petit ou egal a 4")

```

3. *Indentation* :  Python utilise l’indentation pour comprendre la fin des boucles, conditions, fonctions...   
Tester le code suivant, et corriger l’erreur.

<p style="background:#F3F3F3;margin-left: -0.5em;" ><code style="background:#F3F3F3;"># Code à recopier
for i in range(10):
if not i==4:
print(i)</code> </p>

<!-- grade_id: cell-a06c258a39da0fd1, points: 0, grade: false, solution: true -->
```python
# Question 3 avec code indenté

for i in range(10):
    if not i==4:
        print(i)


```

4. *Aide*: Lorsque vous cherchez comment utiliser une fonction, penser à chercher dans l’aide et sur
internet.

```python
s = sum(1,2,3,4) #soulève une erreur à cause des arguments de la fonction sum
```

```python
help(sum)  # il lui faut un "iterable" par exemple une liste
```

```python
# Code corrigé
s = sum([1,2,3,4])
print(s)
```

5. *Variables aléatoires*: Comprendre et expliquer le code suivant (utiliser toutes les ressources d’aide
nécessaire pour comprendre ce que font les fonctions) :

<p style="background:#F3F3F3">
<code style="background:#F3F3F3;"># Code à recopier 1
b = np.random.binomial(1,1/2,30)
print(b) </code> </p>

<p style="background:#F3F3F3">
<code style="background:#F3F3F3;"># Code à recopier 2
b = np.random.binomial(1,1/2,30)
print(b)
# Pourquoi est-ce différent de la première fois ?</code> </p>

<p style="background:#F3F3F3">
<code style="background:#F3F3F3;"># Code à recopier 3
m = np.random.binomial(1,1/2,(4,3))
print(m) </code> </p>

<!-- grade_id: cell-2341250c4068eb5a, points: 0, grade: false, solution: true -->
```python
# Question 5

# Pour utiliser la bibliothèque NumPy il est généralement nécessaire de l'importer :
import numpy as np

# Code à recopier 1
print('Code 1 \n')
b = np.random.binomial(1,1/2,30)
print(b) 
# Code à recopier 2
print('\nCode 2\n')
b = np.random.binomial(1,1/2,30)
print(b)
# Le code génère un tirage d'une variable aléatoire on n'obtient pas deux fois 
# le même tirage d'affiler.

# Code à recopier 3
print('\nCode 3\n')
m = np.random.binomial(1,1/2,(4,3))
print(m) 

```

## Exercice 2 : NumPy

1. Commandes de base

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier
b = np.cos(1); print("b = %s" % b)
bb = np.sqrt(4); print("bb = %s" % bb) 
 </code> </p>

<!-- grade_id: cell-3d20e3a4ae5f7f50, points: 0, grade: false, solution: true -->
```python
# Question 1


b = np.cos(1); print("b = %s" % b)
bb = np.sqrt(4); print("bb = %s" % bb) 

```

2. Création de tableaux NumPy

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 1
u = [1, 9, -4, 0.5]
print(2*u)
</code> </p>

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 2
v = np.array([1, 9, -4, 0.5])
print(2*v)
print(20*"-")
print(v[2])
print(20*"-")
print(v[0])
</code> </p>

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 3
print(v.size)
print(20*"-")
print(np.size(v))
print(20*"-")
print(v.shape)
# la numérotation commence à zero
</code> </p>

<!-- grade_id: cell-2b21537e87b7c0e8, points: 0, grade: false, solution: true -->
```python
# Question 2 


# Code à recopier 1
print('Code 1\n')
u = [1, 9, -4, 0.5]
print(2*u)

# Code à recopier 2
print('\nCode 2\n')
v = np.array([1, 9, -4, 0.5])
print(2*v)
print(20*"-")
print(v[2])
print(20*"-")
print(v[0])

# Code à recopier 3
print('\nCode 3\n')
print(v.size)
print(20*"-")
print(np.size(v))
print(20*"-")
print(v.shape)
# la numérotation commence à zero

```

3. Intervalles et répétitions

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 1
print(np.linspace(0, 2 * np.pi, 10))
print(20*"-")
print(np.linspace(20, 1, 10))
print(20*"-")
x,dx=np.linspace(0,6.3,10,retstep=True)
print(x)
print(20*"-")
print(dx)
</code> </p>

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 2
x = np.arange(0, 10, 0.5)
print(x)
</code> </p>

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 3
dx = 0.5
x = np.arange(0, 10+dx/2, dx)
print(x)
</code> </p>

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 4
y = np.arange(10, -dx/2, -dx)
print(y)
</code> </p>

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 5
print(np.repeat(3.2, 4))
print(20*"-")
x = np.array([[1,2],[3,4]])
print(np.repeat(x, 2))
print(20*"-")
print(np.repeat(x, 3, axis=1))
print(20*"-")
print(np.repeat(x, 4, axis=0))
</code> </p>

<!-- grade_id: cell-386729b6f77a5308, points: 0, grade: false, solution: true -->
```python
# Question 4

# Code à recopier 1
print('Code 1\n')
print(np.linspace(0, 2 * np.pi, 10))
print(20*"-")
print(np.linspace(20, 1, 10))
print(20*"-")
x,dx=np.linspace(0,6.3,10,retstep=True)
print(x)
print(20*"-")
print(dx)

# Code à recopier 2
print('\nCode 2\n')
x = np.arange(0, 10, 0.5)
print(x)

# Code à recopier 3
print('\nCode 3\n')
dx = 0.5
x = np.arange(0, 10+dx/2, dx)
print(x)

# Code à recopier 4
print('\nCode 4\n')
y = np.arange(10, -dx/2, -dx)
print(y)

# Code à recopier 5
print('\nCode 5\n')
print(np.repeat(3.2, 4))
print(20*"-")
x = np.array([[1,2],[3,4]])
print(np.repeat(x, 2))
print(20*"-")
print(np.repeat(x, 3, axis=1))
print(20*"-")
print(np.repeat(x, 4, axis=0))


```

4. Opérations sur les vecteurs

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier
A = np.array([1, 2, 3, 6])
B = np.array([0, -4, 9, 4]);
print(np.exp(A))
print(20*"-")
print(A+B)
print(20*"-")
print(A*B)
print(20*"-")
print(A/B)
print(20*"-")
print(2*A)
print(20*"-")
print(A+1)
print(20*"-")
print(A+np.array([1,2,3]))
</code> </p>

<!-- grade_id: cell-e0dc06587f6ac6ef, points: 0, grade: false, solution: true -->
```python
# Question 4
# Code à recopier
A = np.array([1, 2, 3, 6])
B = np.array([0, -4, 9, 4]);
print(20*"-")
print(np.exp(A))
print(20*"-")
print(A+B)
print(20*"-")
print(A*B)
print(20*"-")
print(A/B) # Problème division par zéro (warning l'exécution continue)
print(20*"-")
print(2*A)
print(20*"-")
print(A+1)
print(20*"-")
print(A+np.array([1,2,3])) # Incompatibilité de dimensions (erreur)
```

5. Sélection

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier
print(A==2)
print(20*"-")
print(A[A>2])
</code> </p>

<!-- grade_id: cell-51863cb6238f5e59, points: 0, grade: false, solution: true -->
```python
# Question 5

# Code à recopier
print(20*"-")
print(A==2)
print(20*"-")
print(A[A>2])
```

6. Sommes et sommes cumulées

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier
u = np.arange(1,10,1)
print(np.sum(u))
print(20*"-")
print(np.cumsum(u))
</code> </p>

<!-- grade_id: cell-79f93de3acdba433, points: 0, grade: false, solution: true -->
```python
# Question 5

# Code à recopier
u = np.arange(1,10,1)
print(np.sum(u))
print(20*"-")
print(np.cumsum(u))
```

7. Création et opération sur les matrices     
Repérez les éventuels problèmes:

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier
A.shape = (2,2); print(A)
print(20*"-")
print(A.transpose())
C = B.reshape(2,2).T
print(C); 
print(20*"-")
print(A @ C)
print(20*"-")
print(np.matmul(A,C))
print(20*"-")
print(C[1,])
print(20*"-")
print(C/A[1,])
</code> </p>

<!-- grade_id: cell-7c5477d85fe79959, points: 0, grade: false, solution: true -->
```python
# Question 7

# Code à recopier
A.shape = (2,2); print(A)
print(20*"-")
print(A.transpose())
C = B.reshape(2,2).T
print(C); 
print(20*"-")
print(A @ C)
print(20*"-")
print(np.matmul(A,C))
print(20*"-")
print(C[1,])
print(20*"-")
print(C/A[1,])

```

## Exercice 3 : Matplotlib

Pour tracer le graphique d’une fonction mathématique, on utilise la bibliothèque *matplotlib*, dont
la partie pertinent *pyplot* a été importée sous le
petit nom de *plt*.

```python
import matplotlib.pyplot as plt
```

1. Recopier et exécuter les codes suivants :

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 1
x = np.linspace(0, 2*np.pi, 100)
plt.plot(x, np.sin(x))
plt.show()
</code> </p>

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 2
plt.plot(x, np.cos(x))
plt.plot(x, np.sin(x))
plt.show()
</code> </p>

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier 3
plt.plot([0, 1], [1, 0], "b",linewidth=2,linestyle='--',label='segment')
plt.plot(0.6, 0.9, "r+",label='point')
plt.xlabel("x")
plt.ylabel("y")
plt.title("joli graphe")
plt.legend()
plt.show()
</code> </p>

<!-- grade_id: cell-8b21d51aacfefe70, points: 0, grade: false, solution: true -->
```python
# Question 1

# Code à recopier 1
plt.figure(1)
x = np.linspace(0, 2*np.pi, 100)
plt.plot(x, np.sin(x))
plt.show()

plt.figure(2)
plt.plot(x, np.cos(x))
plt.plot(x, np.sin(x))
plt.show()

plt.figure(3)
plt.plot([0, 1], [1, 0], "b",linewidth=2,linestyle='--',label='segment')
plt.plot(0.6, 0.9, "r+",label='point')
plt.xlabel("x")
plt.ylabel("y")
plt.title("Joli graphe !")
plt.legend()
plt.show()


```

2. Plusieurs graphes dans la même image :

<p style="background:#F3F3F3"> <code style="background:#F3F3F3;"># Code à recopier
fig,ax=plt.subplots(1,2)
ax[0].plot(x, np.sin(x))
ax[1].plot(y, y*np.sin(y))
fig.suptitle("2 jolis graphes")
ax[0].set_title("sinus")
ax[1].set_xlabel("x")
plt.show()
</code> </p>

<!-- grade_id: cell-a74b1778d2636801, points: 0, grade: false, solution: true -->
```python
# Question 2

# Code à recopier
fig,ax=plt.subplots(1,2)
ax[0].plot(x, np.sin(x))
ax[1].plot(y, y*np.sin(y))
fig.suptitle("2 jolis graphes")
ax[0].set_title("sinus")
ax[1].set_xlabel("x")
plt.show()

```

## Exercices supplémentaires

### Exercice 1 :  La factorielle

Ecrire une fonction qui prend un entier et retourne sa factorielle.   
On rappelle que : $n!=1 \times 2 \times ...\times (n-2) \times(n-1)\times n$.

<!-- grade_id: cell-18f506952b7c5981, points: 0, grade: false, solution: true -->
```python
def factorielle(n):
    prod=1
    for i in range(2,n+1):
        prod*=i
    return prod
```

<!-- grade_id: cell-3151acc2bdf5d7bf, points: 1, grade: true, solution: false -->
```python
assert(factorielle(4)==24)
```

### Exercice 2 :  Coefficient Binômial

On rappelle la formule du coefficient binomial :

$C_{n}^k=\binom{n}{k}=\frac{n!}{k!(n - k)!}.$


<!-- grade_id: cell-bb2fb592d4da275b, points: 0, grade: false, solution: true -->
```python
def coeff_binomial(k,n):
    if (k<0) or (n<0):
        raise ValueError("Entiers positifs pour k et n")
    if k>n:
        return 0
    coeff=factorielle(n)/(factorielle(k)*factorielle(n-k))
    return coeff

```

<!-- grade_id: cell-f02cf6a829654e39, points: 1, grade: true, solution: false -->
```python
assert(coeff_binomial(3,6)==20)

```

### Exercice 3 : Triangle de Pascal

Le triangle de Pascal (https://fr.wikipedia.org/wiki/Triangle_de_Pascal), contient les coefficients binomiaux. On peut le lire comme un tableau triangulaire inférieur où la case à l'intersection entre la ligne $n$ et la colonne $k$ correspond au coefficient $\binom{n-1}{k-1}$.  

On se propose dans cet exercice d'écrire une fonction qui prend en entrée un entier $n$ et qui retourne le triangle de Pascal (sous forme de tableau) à $n+1$ lignes.



<!-- grade_id: cell-92721d47684f596f, points: 0, grade: false, solution: true -->
```python
# Indice : utiliser une double boucle for
def triangle_pascal(n):
    tableau = np.zeros((n+1,n+1),int)
    tableau[:,0]=1;tableau[1,1]=1
    for i in range(2,n+1):
        for j in range(1,i+1):
            tableau[i,j]=tableau[i-1,j-1]+tableau[i-1,j]
    return tableau
    '''
    # Autre solution possible
    for i in range(n+1):
        for j in range(i+1):
            tableau[i,j]=coeff_binomial(i,j)
    return tableau
    '''

    
```

<!-- grade_id: cell-ffdb4593163dd940, points: 1, grade: true, solution: false -->
```python
triangle4=triangle_pascal(4)
assert(triangle4[-1,2]==6)

```

# Exercice 4 : Suite de Syracuse 

La suite de Syracuse est une suite récurrente d'entiers naturels définie comme suit :

$$
u_0 = N \in \mathbb{N},\;\\
u_n = 
\begin{cases}
\frac{u_{n-1}}{2},\; si \; u_{n-1}\; pair\\
3u_{n-1}+1, \; sinon.
\end{cases}
$$

On se propose d'écrire une fonction qui prend en entrée le nombre de termes à calculer $n$ et la valeur du premier terme $u_0$ et qui renvoit un vecteur contenant tous ces termes.  

Une fois la fonction écrite, il est proposé de tracer la suite pour plusieurs valeurs (que vous pouvez varier). Que pouvez-vous dire sur cette suite ?   
Pour plus d'information voir : https://fr.wikipedia.org/wiki/Conjecture_de_Syracuse.

<!-- grade_id: cell-fb2df4c96aaad4e3, points: 0, grade: false, solution: true -->
```python
def syracuse(n,u_0):
    U_syr=np.zeros(n+1,int)
    U_syr[0]=u_0
    for i in range(1,n+1):
        if U_syr[i-1]%2 ==0:
            U_syr[i] = U_syr[i-1]/2
        else:
            U_syr[i] = 3*U_syr[i-1]+1
    return U_syr
    
```

<!-- grade_id: cell-9845ece55eafafd6, points: 1, grade: true, solution: false -->
```python
U10 = syracuse(20,11)
assert(U10[-1]==1)
### BEGIN HIDDEN TESTS
U20 = syracuse(20,13)
assert(U20[-1]==2)
### END HIDDEN TESTS
```

<!-- grade_id: cell-a8bf5f157e0545f0, points: 0, grade: false, solution: true -->
```python
U_syr27 = syracuse(100,27)
U_syr12 = syracuse(100,12)

fig,ax=plt.subplots(1,2)
ax[0].plot(U_syr27,'r+')
ax[1].plot(U_syr12,'r+')

ax[0].set_title('Suite de Syracuse avec $u_0$=27')
ax[1].set_title('Suite de Syracuse avec $u_0$=12')
ax[0].set_xlabel("n")
ax[1].set_xlabel("n")
ax[0].set_ylabel("$u_n$")
ax[1].set_ylabel("$u_n$")

fig.tight_layout()

plt.show()

```