```python
# A faire executer une fois au début
! pip install pydataset
```

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import scipy.stats as st
from pydataset import data
```

# 1. La loi des grands nombres

## Exercice 4.1

1. En utilisant <code>st.binom.rvs</code>, simuler un échantillon $(X_1,X_2,\ldots,X_n)$ de variables indépendantes de même loi de Bernoulli $\mathcal{B}(1,0.5)$ pour $N$ grand (par exemple $N=2000$).

<!-- grade_id: cell-b113710092cebd01, points: 0, grade: false, solution: true -->
```python
# Question 1

N=2000
X=
```

2. Illuster la loi des grands nombres en traçant $\bar{X_n} = \frac{X_1+\ldots,X_n}{n}$ en fonction de $n$ (compris entre $1$ et $N$). On tracera aussi la limite attendue sur le même graphique. 

On pourra pour cela utiliser la fonction <code>np.cumsum</code> pour construire
$$ Y_n = X_1+X_2+\ldots+X_n $$

<!-- grade_id: cell-e2a2fda9c0dcd8fe, points: 0, grade: false, solution: true -->
```python
# Question 2

Ybar = 
Xbar =
```

3. Procéder de même avec un échantillon de loi normale $\mathcal{N}(1,4)$ de moyenne 1 et de variance 4.


<!-- grade_id: cell-10eb43f44bb24f49, points: 0, grade: false, solution: true -->
```python
# Question 3

N=2000
X=
Xbar=

```

4. Recommencer avec un échantilon de variables de cauchy $C(0,1)$. Que remarquez-vous ?   
Indication: On produira l'échantillon par la commande
<p style="background:#F3F3F3;" >
<code style="background:#F3F3F3;"> N = 10000; Z = st.cauchy.rvs(size=N,loc=0,scale=1). </code> </p>

<!-- grade_id: cell-c119e1cb0618f1e4, points: 0, grade: false, solution: true -->
```python
# Question 4


```

# 2. Théorème centrale limite

## Exercice 4.2 

1. Soit $S_N$ de loi binomiale $\mathcal{B}(N,p)$ avec $p=0.5$ et $N=1000$.   
On définit $\bar{X_N} = \frac{S_N}{N}$ et la v.a centrée réduite associée :
$$ Z_N = \sqrt{N} \frac{\bar{X}_N-p}{\sqrt{p(1-p)}}$$

On donne la fonction <code>varTCL</code>, que fait-elle ?   
Utiliser cette fonction pour simuler un échantillon de $M = 5000$ variables de même loi que $Z_N$.



<!-- grade_id: cell-dda5e0ad3892f578, points: 0, grade: false, solution: true -->
```python
# Question 1

def varTCL(S):
    return((S/N-p)/np.sqrt(p*(1-p)/N))

N=#...
M=#...
p=#...
S=st.binom.rvs(N,p,size=M)
Z=



```

2. Illustrer le Théorème centrale limite en traçant l'histogramme des $M$ réalisations et en le comaprant à la densité de la loi $\mathcal{N}(0,1)$.   
Utiliser <code> st.norm.pdf </code>.

```python
# Question 2

plt.figure()
plt.hist(Z,bins=20,density=True)

# Loi normale centrée réduite : desnité sur [-3,3]
x=np.linspace(start=-3,stop=3,num=600)
plt.plot(x,"...",color="red");

```

3. On propose ici de compléter le code suivant afin de faire varier $M$ et observer la convergence des $Z_N$ vers une loi normale centrée réduite.

```python
from ipywidgets import interact;
import ipywidgets as widgets;

def Simulation_Z(M,N=10000,p=0.5):
    S="..."
    def varTCL(s):
        return((s/N-p)/np.sqrt(p*(1-p)/N))
    Z=["..." for s in S]
    return(Z)

def plot(M):
    plt.figure(figsize=(10,6));
    Z = Simulation_Z(M);
    plt.hist(Z,bins=25,density=True,alpha=0.4,color="black");
    
    x=np.linspace(start=-3,stop=3,num=600);
    plt.plot(x,"...",color="red",linestyle='--',linewidth=2);
    plt.xlim([-3,3]);
    plt.ylim([0,0.5]);


interact(plot, M = widgets.IntSlider(value=50, min=50, max=2000, step=50));
```

# 3. Intervalles de confiance
## Exercice 4.3
Jean souhaite connaître la taille moyenne des garçons de sa promo. Cependant il y a des milliers d'élèves dans sa promi et il ne peut pas tous les interroger un à un. A la sortie de l'amphi de Statistiques, il a pu récolter la taille de plusieurs de ses camarades (en cm):
$$ 180,170,186,184,171,184,167,180,177.$$

1. Quelles hypothèses doit-on admettre pour calculer un intervalle de confiance de la taille moyenne ?

2. On sait que la variance $\sigma^2$ est égale à 5. Calculer un intervalle de confiance (bilatéral) de niveau $95\%$ de la taille moyenne.

<!-- grade_id: cell-9405c21b06b88401, points: 0, grade: false, solution: true -->
```python
# Question 2

taille=[180, 170, 186, 184, 182, 171, 184, 167, 180, 177]
n = "..."
moyenne = "..."
sigma = "..."
delta_95 = "..."
IC_95 = (moyenne - delta_95, moyenne + delta_95)
print("IC au niveau 0.95 :[%.4f,%.4f]" % IC_95)
```

3. Jean a réussi à récolter 10 nouvelles tailles:
$$ 173,167,170,174,178,178,175,166,172,170.$$
Comment évolue l'intervalle de confiance (bilatéral) de niveau $95\%$

<!-- grade_id: cell-c63efe5b9c628df1, points: 0, grade: false, solution: true -->
```python
# Question 3

taille = np.concatenate((taille,[173, 167, 170, 174, 178, 178, 175, 166, 172, 170]))
n = "..."
moyenne = "..."
sigma = "..."
delta_95 = "..."
IC_95 = (moyenne - delta_95, moyenne + delta_95)
print("IC au niveau 0.95 :[%.2f,%.2f]" % IC_95)

```

## Exercice 4.4
Suite aux travaux menés par Jean, Amina souhaite faire la même chose pour les filles de sa promo. Cependant elle ne connaît pas la
variance de la taille des filles. Voici les tailles collectées par Amina :
$$ 166,170,161,167,168,169,169,166,163,161,162,171,169,156,168.$$

1. En utilisant la fonction <code>st.t.interval</code>, (aide en ligne, on pourra utiliser <code>st.sem</code> pour l'échelle), donner un intervalle de confiance (bilatéral) de niveau $95\%$ de la taille moyenne des filles.

<!-- grade_id: cell-140816f913facfba, points: 0, grade: false, solution: true -->
```python
# Question 1

taille=[ 166, 170, 161, 167, 168, 169, 169,166, 163, 161, 162, 171, 169, 156, 168]

```

2. Ecrire une fonction prenant en argument un échantillon et un niveau de confiance et retournant l'intervalle de confiance associé pour la moyenne.

<!-- grade_id: cell-118a39f1c2421c89, points: 0, grade: false, solution: true -->
```python
# Question 2 
def StudentInterval(x,alpha=0.05):

    m = "moyenne à compléter"
    s = "écart type non-biasé à compléter"
    l = "Taille échantillon à compléter"
    delta = (st.t.ppf(1-alpha/2,df=l-1)*s) / np.sqrt(l)
    return (m-delta,m+delta)

```

<!-- grade_id: cell-9c029774974fe744, points: 1, grade: true, solution: false -->
```python
assert(st.t.interval(0.95,df=len(taille)-1,loc=np.mean(taille),scale=st.sem(taille)) == StudentInterval(taille))
```

```python
StudentInterval(taille,alpha=0.1)
```

# 4. Tests d’hypothèses et prédiction à partir d’une Régression linéaire

On rappelle que dans un test d’hypothèse, on cherche à tester une hypothèse nulle $H_0$ contre une alternative $H_1$. Certaines fonctions de <code>scipy.stat</code> permettent de calculer la grandeur quantitative utilisée pour ce test, nommé $p$-valeur (cf. CM). On rappelle les lignes directrices suivantes pour conclure le test à partir de la $p$-valeur:
- $p \leq 0.01$: très forte présomption contre $H_0$
- $0.01 < p \leq 0.05$: forte présomption contre $H_0$ (donc rejet de $H_0)$
- $0.05 <p < \leq 0.1$: faible présomption contre $H_0$ (mais pas de rejet de $H_0$)
- $p>0.1$: pas de présomption contre $H_0$

## Exercice 4.5 
On reprend le jeu de données du TP2 (exercice 6)

### Importer les données avec Google Colab 

Importer le fichier "Donnes.csv". Pour cela placer le fichier dans le dossier de votre drive où est situé ce notebook, puis recopier et exécuter les cellules suivantes.

<code>from google.colab import drive </code>

<code>import os</code>

<code>drive.mount('/content/drive')</code>

```python
# Compléter /... avec le path vers le dossier où sont situés le Notebook et le fichier .csv
# Décommenter les trois lignes :
# - os.getcwd() permet de voir votre path courant
# - os.listdir() liste les fichiers et dossiers de votre path courant.
# Utiliser ses lignes en cas d'erreur pour voir si le path est bon.

### 3 LIGNES A DECOMMENTER  ###
#os.chdir("drive/My Drive/Colab Notebooks/...")
#print(os.getcwd())
#os.listdir()
```

### Importer les données avec Jupyter Notebook

Placer le fichier "Donnees.csv"  dans le même dossier que le ipynb. Importer ce fichier avec la function <code>pd.read_csv("fichier.csv",sep="\t") </code>.

### Exercice 

Les données correspondent à l’âge, au poids, à la taille, à la consommation hebdomadaire d’alcool (en nombre de verres bus), au sexe, au ronflement et au tabagisme, d’un échantillon de 100 personnes.

1. Créer deux variables <code>th</code> et <code>tf</code> regroupant les tailles des hommes et des femmes dans cet échantillon. Quelles hypothèses doit on faire pour effectuer un test de Student sur les moyennes $\mu_h$ de <code>th</code> et $\mu_f$ de <code>tf</code>

```python
# Question 1

Donnees= pd.read_csv("Donnees.csv",sep="\t") 

# Création des séries pour chaque variable
for nom in Donnees.keys():
    globals()[nom] = Donnees[nom]

th = 
tf =

```

2. En utilisant la fonction <code>st.ttest_1samp</code>, testez si la moyenne de l’échantillon de la taille des hommes diffère significativement de la moyenne française, autrement dit, effectuez le test de l’hypothèse nulle $H_0 : \mu_h = 176,6$ contre l’hypothèse alternative : $H_1: \mu_h \neq 176,6$.

```python
# Question 2

print(st.ttest_1samp(th, popmean = "COMPLETER"))
```

3. Tester de même si la moyenne de la taille des femmes de l’échantillon dépasse significativement la moyenne française 163,9 cm. Dépasse-t-elle significativement 176,6 ?

```python
# Question 3

print(st.ttest_1samp(tf, popmean = "COMPLETER", alternative = "COMPLETER"))

print(st.ttest_1samp(th, popmean = "COMPLETER", alternative = "COMPLETER"))
```

4. Calculer la table de contingence entre les variables <code>tabac</code> et <code> ronfle</code> (rappel : utiliser <code> pd.crosstab </code>).   
En lui appliquant la <code>st.chi2_contingency</code>, tester si tabagisme et ronflement sont indépendants (test d’indépendance du $\chi^2$).

```python
# Question 4

table = "COMPLETER"
print(table,'\n')
print(st.chi2_contingency(table))
```

5. Réaliser de même un test d’indépendance entre consommation d’alcool et sexe.

**Indication:** pour respecter la condition sur l'effectif des populations, on commencera par créer une variable ordinale pour la consommation d'alcool avec 3 valeurs «Aucun verre»,
«Entre 1 à 4 verres» et «Plus de 5 verres».

```python
# Question 5

table = "COMPLETER"
print(table,'\n')
print(st.chi2_contingency(table))
# Quel est le problème

```

```python
# Question 5 avec indice 

def RegroupeAlcool(x):
    if "COMPLETER": return("COMPLETER")
    elif "COMPLETER": return("COMPLETER")
    else: return("COMPLETER")

Alcool=alcool.apply(RegroupeAlcool)
Alcool=Alcool.astype('category')
```

## Exercice 4.6 
On reprend un jeu de données du TP1, qu’on charge en exécutant les commandes suivantes:

```python
from pydataset import data
iris=data("iris");
y=iris["Petal.Length"]
x=iris["Sepal.Length"]
l=len(x)
```

1. En supposant que $x$ et $y$ sont des échantillons normaux, tester (avec <code>st.pearsonr</code>) si ils sont significativement corrélés.

<!-- grade_id: cell-cbec7202d2c8e9c4, points: 0, grade: false, solution: true -->
```python
# Question 1


```

2. On cherche à étudier la régression linéaire de la longueur des Pétales en fonction de la longueur des Sépales. Quelles informations tirez vous des commandes suivantes ?
<p style="background:#F3F3F3;" >
<code style="background:#F3F3F3;">import statsmodels.api as sm
X=np.column_stack((x,np.ones(l)))
res=sm.OLS(y,X).fit()
print(res.summary())</code> </p>

```python
# Question 2

```

3. Pour obtenir un intervalle de prédiction au niveau $95\%$, on utilise la commande:
    <p style="background:#F3F3F3;" >
<code style="background:#F3F3F3;">prediction=res.get_prediction().summary_frame(alpha=0.05)</code> </p>
Tracer sur un même dessin un nuage de points de la variable $y$ en fonction de la variable $x$, la droite de régression en bleue et en rouge les deux droites bornant inférieurement (<code>prediction["obs_ci_lower"]</code>) et supérieurement( <code>prediction["obs_ci_upper"]</code>) les intervalles de prédiction au niveau de confiance $95\%$.

<!-- grade_id: cell-a8ed39e449b33606, points: 0, grade: false, solution: true -->
```python
# Question 3

```

4. Dans quel intervalle prédisez-vous que se trouvera la longueur des Pétales si la longueur des Sépales est $x = 7.2$ ? (avec un niveau de confiance $95\%$)  
Réutiliser <code> res.get_prediction("COMPLETER").summary_frame(alpha=0.05) </code> pour effectuer une prédiction.

<!-- grade_id: cell-a399914d65122b47, points: 0, grade: false, solution: true -->
```python
# Question 4

```